```python
# Initialize Otter
import otter
grader = otter.Notebook("hw05.ipynb")
```

# Homework 5A: Sampling
## Due Date: Thursday, February 29th, 11:59 PM

You must submit this assignment to Gradescope by the on-time deadline, Thursday, February 29th, 11:59 PM. Please read the syllabus for the grace period policy. No late submissions beyond the grace period will be accepted. While course staff is happy to help you if you encounter difficulties with submission, we may not be able to respond to late-night requests for assistance (TAs need to sleep, after all!). **We strongly encourage you to plan to submit your work to Gradescope several hours before the stated deadline.** This way, you will have ample time to contact staff for submission support. 

This is part of a two-part assignment. After completing this part ("Homework 5A"), please submit the generated zip file to the Homework 5A Coding assignment on Gradescope. Gradescope will automatically submit a PDF of your written responses to the HW 5A Written assignment; there is no need to submit it manually. **Note that you will have to submit HW 5B manually and assign pages accordingly.**

## Collaboration Policy

Data science is a collaborative activity. While you may talk with others about the homework, we ask that you **write your solutions individually**. If you do discuss the assignments with others please **include their names** below.

**Collaborators**: *list collaborators here*

## This Assignment

The purpose of this assignment is for you to understand the process of sampling: including convenience samples, random samples; as well as sampling biases. We will also practice drawing samples from a categorical distribution using `NumPy`.

For the purposes of clarity, we will refer to all variations of this class (Data 100/200/200S) as Data 100 in this homework. 

## Score Breakdown
Question | Manual? | Points
--- | --- | --
1ai | No | 1
1aii | No | 1
1aiii | No | 1
1bi | No | 1
1bii | No | 1
1biii | No | 1
1biv | No | 1
2a | No | 2
2b | Yes | 2
2c | Yes | 3
2d | Yes | 2
Total | 3 | 16


## Before You Start

For each question in the assignment, please write down your answer in the answer cell(s) right below the question.

We understand that it is helpful to have extra cells breaking down the process towards reaching your final answer. If you happen to create new cells below your answer to run code, **NEVER** add cells between a question cell and the answer cell below it. It will cause errors when we run the autograder, and it will sometimes cause a failure to generate the PDF file.

**Important note: The local autograder tests will not be comprehensive. You can pass the automated tests in your notebook but still fail tests on Gradescope after the grades are released.** Please be sure to check your results carefully.

Finally, unless we state otherwise, **do not use for loops or list comprehensions**. The majority of this assignment can be done using built-in commands in `pandas` and `NumPy`.

### Debugging Guide

If you run into any technical issues, we highly recommend checking out the [Data 100 Debugging Guide](https://ds100.org/debugging-guide/). In this guide, you can find general questions about Datahub, Gradescope, and common `pandas` and RegEx errors.

### Initialize your environment

```python
import numpy as np
np.random.seed(200)
import matplotlib
import matplotlib.pyplot as plt
import seaborn as sns
plt.style.use('fivethirtyeight')
```

<br/><br/>
<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Data 100 Cutest Pets Contest

Welcome to the Data 100 Cutest Pets Contest, Spring 2024 edition! Course staff nominate their pets to participate in this contest. Students will vote on the cutest one among the nominations in the final exam.

The nominees are:

<img src='images/appa.jpg' width="180px" />  <img src='images/pishi.png' width="180px" /> <img src='images/mimi.png' width="180px" /> 

(From left to right) **Appa** (Matthew's cat), **Pishi** (Professor Norouzi's cat), and **Mimi** (Shiny's dog).

Course staff would like to predict the results for the official survey later in the semester by surveying students in the class now. This process is similar to polling that occurs before a political election.

In this section, you are going to explore different sampling methods.
<br/><br/>

<hr style="border: 1px solid #fdb515;" />

## Question 1

<br>

---

### Question 1a

Since her dog, Mimi, is nominated, Shiny would like to understand the class opinion before the contest. This coming week, she decided to survey all students enrolled in Data 100 this Spring semester (Spring 2024) by sending out an Ed announcement via email that asked students to choose the cutest from the three pets. You may assume no other students/users receive the survey. Shiny closes the survey 12 hours after sending it out.

You can assume that all, and only, enrolled students are on Ed.

#### Question 1a, Part i
In Shiny's survey, which of the following is the population of interest? Assign your answer choice to `q1ai` (e.g., `'A'`, `'B'`, etc.).

A. All UC Berkeley students<br/>
B. All students enrolled in Data 100 across all semesters (Spring 2024 and previous)</br>
C. All students enrolled in Data 100 for this semester (Spring 2024)<br/>
D. All students who fill out Shiny's survey

```python
q1ai = ...
```

```python
grader.check("q1ai")
```

#### Question 1a, Part ii
In Shiny's survey, which of the following is the sampling frame? Assign your answer choice to `q1aii` (e.g., `'A'`, `'B'`, etc.).

A. UC Berkeley students<br/>
B. All students enrolled in Data 100 across all semesters (Spring 2024 and previous)</br>
C. All students enrolled in Data 100 for this semester (Spring 2024)<br/>
D. All students who fill out Shiny's survey

```python
q1aii = ...
```

```python
grader.check("q1aii")
```

#### Question 1a, Part iii

Which of the following is the sample? Assign your answer choice to `q1aiii` (e.g., `'A'`, `'B'`, etc.).

A. UC Berkeley students<br/>
B. All students enrolled in Data 100 across all semesters (Spring 2024 and previous)</br>
C. All students enrolled in Data 100 for this semester (Spring 2024)<br/>
D. All students who fill out Shiny's survey

```python
q1aiii = ...
```

```python
grader.check("q1aiii")
```

<br><br>

---

### Question 1b

In practice, we cannot get a 100\% survey response rate, often because our population is too large, or because there is a time limit. In this case, very few students answered Shiny's survey before she closed it.

To get more data to predict the answer to the original question ("Which pet will win the Data 100 Cutest Pet Contest?"), Shiny decides on a different strategy: **she conducts the pre-contest survey in person in her discussion section that same week**. She then asks every student who attends the discussion that week for their opinion on the cutest of the three pets, by presenting the following slide: 

<img src='images/survey.png' width="600px" />

#### Question 1b, Part i
In this sampling scheme, which of the following is the population of interest? Assign your answer choice to `q1bi` (e.g., `'A'`, `'B'`, etc.).

A. UC Berkeley students<br/>
B. All students enrolled in Data 100 across all semesters (Spring 2024 and previous)</br>
C. All students enrolled in Data 100 for this semester (Spring 2024)<br/>
D. All students enrolled in Shiny's discussion section<br/>
E. All students who fill out Shiny's pre-contest survey

```python
q1bi = ...
```

```python
grader.check("q1bi")
```

#### Question 1b, Part ii
In this sampling scheme, which of the following is the sampling frame? Assign your answer choice to `q1bii` (e.g., `'A'`, `'B'`, etc.).

A. UC Berkeley students<br/>
B. All students enrolled in Data 100 across all semesters (Spring 2024 and previous)</br>
C. All students enrolled in Data 100 for this semester (Spring 2024)<br/>
D. All students enrolled in Shiny's discussion section<br/>
E. All students who fill out Shiny's pre-contest survey

```python
q1bii = ...
```

```python
grader.check("q1bii")
```

#### Question 1b, Part iii
Which of the following is the sample? Assign your answer choice to `q1biii` (e.g., `'A'`, `'B'`, etc.).

A. UC Berkeley students<br/>
B. All students enrolled in Data 100 across all semesters (Spring 2024 and previous)</br>
C. All students enrolled in Data 100 for this semester (Spring 2024)<br/>
D. All students enrolled in Shiny's discussion section<br/>
E. All students who fill out Shiny's pre-contest survey

```python
q1biii = ...
```

```python
grader.check("q1biii")
```

#### Question 1b, Part iv
Which of the following best characterizes the sample? Assign your answer choice to `q1biv`  (e.g., `'A'`, `'B'`, etc.)..

A. Simple Random Sample<br/>
B. Convenience Sample<br/>
C. Probability Sample

```python
q1biv = ...
```

```python
grader.check("q1biv")
```

<br>

<hr style="border: 1px solid #fdb515;" />


## Question 2
Shiny was able to sample 50 students in her section; the results are as follows:

Pet | Vote Share
--- | ---
Appa | 2%
Pishi | 22%
Mimi | 76%

Based on this result, she predicts that her dog Mimi will win the contest.

Fast-forward to the end of the semester, when the contest has actually taken place. Assume that after tallying every student's votes, **the true popularity of each pet** is: 

Dog | Vote Share
--- | ---
Appa | 12%
Pishi | 80%
Mimi | 8%

The true winner was actually Pishi! Shiny was devastated: what went wrong?

---

### Question 2a

Pishi ultimately wins the Cutest Pet Contest, but from Shiny's results, it was predicted that Mimi would win. Perhaps there was some underlying issue with Shiny's sampling method.

Probability samples can help us quantify sampling bias and chance error. Put briefly, if we assume that a sample distribution was selected at random from a known population, then we can quantify how likely that sample is to have arisen due to random chance (**chance error**). If the difference in sample and population distributions is too great, then we suspect that the given sample has **bias** in how it was selected from the population.

Write one line of code that runs 1000 independent simulations, where each simulation finds the proportion of voters who voted for **Pishi** in a sample of size 50, selected uniformly at random from the **true population**. You may assume that the true population is large enough such that the sample is a random sample with replacement. The output `samples` should be an array with 1000 elements, each of which is the proportion of **Pishi** votes in that simulated sample.

**Hint**: 

* Use `np.random.multinomial` ([documentation](https://numpy.org/doc/stable/reference/random/generated/numpy.random.multinomial.html)).
* Use `array[:, i]` to select the i-th column of a 2D `NumPy` array
* Feel free to print out `samples` or `samples.shape` to check the output!

```python
samples = ...
samples[:5]
```

```python
grader.check("q2a")
```

<!-- BEGIN QUESTION -->

<br>

---

### Question 2b

Create a plot using any `seaborn` and/or `matplotlib.pyplot` functions of your choice to visualize `samples`, which is the simulated distribution of Pishi votes using a sample of size 50. Include descriptive titles and labels. An example is included below. The total area under the plot must be normalized to 1. Your plot may not match exactly ours due to randomness of the data generating process in `np.random.multinomial`.

<img src='images/q2b.png' width="500px" />


**Hint**: use `plt.xlim(left, right)` [(documentation)](https://matplotlib.org/stable/api/_as_gen/matplotlib.pyplot.xlim.html) to specify the left and right limits of the x-axis.

```python
...
plt.show()
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

<br>

---

### Question 2c

According to Shiny's 50-person sample, 22\% of her discussion section reported that they would vote for Pishi in the end-of-semester contest.

In the cell below, create a plot using any `seaborn` and/or `matplotlib.pyplot` functions of your choice to visualize Shiny's sample statistic superimposed on the simulated sample distribution you plotted in the previous part. In other words, include
- a vertical line that passes through 22\%, 
- a vertical line that passes through the mean of the simulated sample distribution, and 
- the simulated sample distribution itself. 

You should choose contrasting colors and include a descriptive title, labels, and a legend if needed. An example is included below.

<img src='images/q2c.png' width="500px" />

```python
...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

<br>

---

### Question 2d

Based on your analysis above, could Shiny's result have arisen due to chance alone? If not, what could be a potential source of bias?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<br/><br/>
<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Congratulations! You have finished Homework 5A!

### Course Content Feedback

If you have any feedback about this assignment or about any of our other weekly, weekly assignments, lectures, or discussions, please fill out the [Course Content Feedback Form](https://docs.google.com/forms/d/e/1FAIpQLSe0fBEJwt6aEfZxU3fh3llNk8rSWHj6Umq0km3wPqmFu0MlGA/viewform). Your input is valuable in helping us improve the quality and relevance of our content to better meet your needs and expectations!

### Submission Instructions

Below, you will see one cell. Running this will automatically generate a zip file with your answers. Please submit this file to the Homework 5A Coding assignment on Gradescope. Gradescope will automatically submit the PDF from this file to the Homework 5A Written assignment. **There is no need to manually submit Homework 5A Written answers; however, please check that the PDF was generated and submitted correctly**. If you run into any issues when running this cell, feel free to check this [section](https://ds100.org/debugging-guide/autograder_gradescope/autograder_gradescope.html#why-does-grader.exportrun_teststrue-fail-if-all-previous-tests-passed) in the Data 100 Debugging Guide.

You are responsible for submitting the zip file generated below as well as a PDF with your answers to Homework 5B. Here are the different components for this homework and the respective submission processes:
* **Homework 5A Coding**: Submit your Jupyter Notebook zip file for Homework 5A,
which can be generated and downloaded from DataHub by using the `grader.export()`
cell provided below.
* **Homework 5A Written**: Gradescope will automatically submit the PDF from the zip file submitted earlier. You do not need to submit anything to this assignment yourself, but please check that the submission went through properly.  
* **Homework 5B**: Submit a PDF to Gradescope that contains your answers to all questions in Homework 5B.

**Important**: When submitting Homework 5B on Gradescope, you **must tag pages to each question correctly** (it prompts you to do this after submitting your work). This significantly streamlines the grading process for our readers. Failure to do this may result in a score of 0 for untagged questions.

To receive credit on this assignment, **you must submit both your coding and written
portions to their respective Gradescope portals.** 

**You are responsible for ensuring your submission follows our requirements and that the automatic submission for Homework 5A Written answers went through properly. We will not be granting regrade requests nor extensions to submissions that don't follow instructions.** If you encounter any difficulties with submission, please don't hesitate to reach out to staff prior to the deadline.

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(run_tests=True)
```

Tests for q1ai (points: 1)
```
>>> assert q1ai.upper() in ['A', 'B', 'C', 'D']
```


Tests for q1aii (points: 1)
```
>>> assert q1aii.upper() in ['A', 'B', 'C', 'D']
```


Tests for q1aiii (points: 1)
```
>>> assert q1aiii.upper() in ['A', 'B', 'C', 'D']
```


Tests for q1bi (points: 1)
```
>>> assert q1bi.upper() in ['A', 'B', 'C', 'D', 'E']
```


Tests for q1bii (points: 1)
```
>>> assert q1bii.upper() in ['A', 'B', 'C', 'D', 'E']
```


Tests for q1biii (points: 1)
```
>>> assert q1biii.upper() in ['A', 'B', 'C', 'D', 'E']
```


Tests for q1biv (points: 1)
```
>>> assert q1biv.upper() in ['A', 'B', 'C', 'D', 'E']
```


Tests for q2a (points: 2)
```
>>> assert all(samples <= 1)
>>> assert all(samples >= 0)
>>> assert samples.shape == (1000,)
```
